<div align="center">

# <b style="font-family: 'LUISS', 'Lato'">Python Exercises - Part IX</b>
<h2 style="font-family: 'LUISS', 'Lato'">Python and R for Data Science</h2>

<h3 style="font-family: 'LUISS', 'Lato'">Data Science and Management</h3>
<img src="https://lollo93.gitlab.io/python-and-r-labs/dist/img/cliente-luiss.png">

<br><br><br>

</div>

---

#### Exercise 1

In the K-Means algorithm, data points are partitioned into clusters based on their proximity to predefined cluster centroids. Each point belongs to the cluster of the nearest centroid. This exercise aims to compute cluster assignments for multiple points based on their distances to the centroids.

Write a function `assign_clusters(points, centroids)` that:
   - Takes a list of 2D points, where each point is a list or tuple of coordinates, e.g., `[[x1, y1], [x2, y2], ...]`.
   - Takes a list of cluster centroids, where each centroid is also a list or tuple of coordinates.
   - Returns a list of integers where each integer represents the index (0-based) of the cluster to which a point belongs. To this aim,
        - For each point in the input, compute the **Euclidean distance** to all centroids:
          
           $\sqrt{(x_2 - x_1)^2 + (y_2 - y_1)^2}$
        
        - Assign each point to the centroid with the smallest distance and include the index of that centroid in the output list.

**NOTES**:
- Use Python's built-in functions or libraries (`math` for square root) for the calculations.
- Do not use pre-built clustering libraries like `sklearn` for this task.

Write your code here:

```python
### Write your code here
import math
def assign_clusters(points, centroids):
    result = []
    for point in points:
        min_dist = math.sqrt((centroids[0][0] - point[0])**2 + (centroids[0][1] - point[1])**2 )
        current_centroid = 0
        for i in range(1, len(centroids)):
            dist = math.sqrt((centroids[i][0] - point[0])**2 + (centroids[i][1] - point[1])**2 )
            if dist < min_dist:
                min_dist = dist
                current_centroid = i
        result.append(current_centroid)
    return result
        

```

To test your code run these cells:

<!-- grade_id: correct_squares_1, points: 1, grade: true, solution: false -->
```python
points = []
centroids = [[1.0, 1.0], [4.0, 4.0]]
expected = []
result = assign_clusters(points, centroids)
assert result == expected, f"Test Failed: Obtained {result}, Expected {expected}"
```

<!-- grade_id: correct_squares_2, points: 2, grade: true, solution: false -->
```python
points = [[2.0, 2.0]]
centroids = [[3.0, 3.0]]
expected = [0]
result = assign_clusters(points, centroids)
assert result == expected, f"Test Failed: Obtained {result}, Expected {expected}"
```

<!-- grade_id: correct_squares_3, points: 3, grade: true, solution: false -->
```python
points = [[2.5, 3.0], [5.0, 1.0], [4.0, 5.0]]
centroids = [[1.0, 1.0], [4.0, 4.0], [6.0, 2.0]]
expected = [1, 2, 1]
result = assign_clusters(points, centroids)
assert result == expected, f"Test Failed: Obtained {result}, Expected {expected}"
```

---
### Exercise 2

Create a class `RectangleManager` to manage `Rectangle`s with sides parallel to the x and y axes.

The class `Rectangle` should have:
- the following attributes:  
    - Bottom-left vertex (`bs`): A tuple representing its coordinates in the plane.  
    - Top-right vertex (`ad`): A tuple representing its coordinates in the plane.
- the following methods:
    - `__init__(self, bs, ad)`: initialize the rectangle, raising a `ValueError` exception when `bs` is greater than `ad`. Notice `ValueError` may take one argument (the message describing the error).
    - `get_area(self)`: returns the area of the rectangle. Notice that:
        - $\text{width} = \texttt{ad[0]} - \texttt{bs[0]}$
        - $\text{height} = \texttt{ad[1]} - \texttt{bs[1]}$
    - `get_perimeter(self)`: returns the perimeter of the rectangle
    - `__str__(self)`: returns a string representing the rectangle in the format `"[(1,3), (7,9)]"` with `(1,3)` being `bs` and `(7,9)` being `ad`

For instance, given a rectangle with:
- Bottom-left vertex (`bs`): (2, 3)  
- Top-right vertex (`ad`): (6, 9)
  
Then we should get:
- `get_area()`: `24`
- `get_perimeter()`: `20`
- `__str__()`: `"[(2,3), (6,9)]"`

The class `RectangleManager` should have:
- a list `rectangles` of rectangles as attribute 
- the following methods:
    - `__init__(self)`: Initialize the class. 
    - `add_rectangle(self, bs, ad)`: Reads the coordinates of the two vertices (`bs` and `ad`) to insert a new rectangle.
    - `delete_rectangles_area_gt(self, area)`: Removes all rectangles with an area greater than a given input value.
    - `__str__(self)`: Returns a multiline string representing for each line a rectangle. Each line contains, separated by one space, the following info:
        - return value `__str__` of the rectangle
        - the area of the rectangle
        - the perimeter of the rectangle<br>
        
    E.g., The rectangle `[(1, 1), (4, 5)]` will be represented as `"[(1, 1), (4, 5)] 12 14"`

Write your code here:

```python
### Write your code here
class Rectangle():
    def __init__(self, bs, ad):
        if bs > ad:
            raise(ValueError("Invalid input"))
        self.bs = bs
        self.ad = ad
    def get_area(self):
        return (self.ad[0] - self.bs[0])*(self.ad[1] - self.bs[1])
    def get_perimeter(self):
        return ((self.ad[0] - self.bs[0]) + (self.ad[1] - self.bs[1])) * 2
    def __str__(self):
        return f"[{self.bs}, {self.ad}]"

class RectangleManager():
    def __init__(self):
        self.rectangles = []
    def add_rectangle(self, bs, ad):
        self.rectangles.append(Rectangle(bs, ad))
    def delete_rectangles_area_gt(self, area):
        result = []
        for rectangle in self.rectangles:
            if rectangle.get_area() <= area:
                result.append(rectangle)
        self.rectangles = result
    def __str__(self):
        result = f"[{self.rectangles[0].bs}, {self.rectangles[0].ad}] {self.rectangles[0].get_area()} {self.rectangles[0].get_perimeter()}"
        for i in range(1, len(self.rectangles)):
            result += f"\n[{self.rectangles[i].bs}, {self.rectangles[i].ad}] {self.rectangles[i].get_area()} {self.rectangles[i].get_perimeter()}"
        return result


```

To test your code run these cells:

<!-- grade_id: correct_sum_of_squares_1, points: 1, grade: true, solution: false -->
```python
rect = Rectangle((2, 3), (6, 9))
assert rect.get_area() == 24, f"Test 1a Failed: Area = {rect.get_area()}, Expected = 24"
```

<!-- grade_id: correct_sum_of_squares_2, points: 1, grade: true, solution: false -->
```python
rect = Rectangle((2, 3), (6, 9))
assert rect.get_perimeter() == 20, f"Test 1b Failed: Perimeter = {rect.get_perimeter()}, Expected = 20"
```

<!-- grade_id: correct_sum_of_squares_3, points: 1, grade: true, solution: false -->
```python
rect = Rectangle((2, 3), (6, 9))
assert str(rect) == "[(2, 3), (6, 9)]", f"Test 1c Failed: String = {str(rect)}, Expected = '[(2, 3), (6, 9)]'"
```

<!-- grade_id: cell-d6ffad607e1095f9, points: 1, grade: true, solution: false -->
```python
manager = RectangleManager()
manager.add_rectangle((1, 1), (4, 5))
manager.add_rectangle((0, 0), (2, 2))
manager.add_rectangle((3, 3), (5, 7))
expected_str = [
    "[(1, 1), (4, 5)]",
    "[(0, 0), (2, 2)]",
    "[(3, 3), (5, 7)]"
]
assert [str(r) for r in manager.rectangles] == expected_str, f"Test 2 Failed:\nGot: {str([str(r) for r in manager.rectangles])}\nExpected: {expected_str}"
```

<!-- grade_id: cell-747703e2b65a4e13, points: 1, grade: true, solution: false -->
```python
manager = RectangleManager()
manager.add_rectangle((1, 1), (4, 5))
manager.add_rectangle((0, 0), (2, 2))
manager.add_rectangle((3, 3), (5, 7))
manager.delete_rectangles_area_gt(8)
expected_str = [
    "[(0, 0), (2, 2)]",
    "[(3, 3), (5, 7)]"
]
assert [str(r) for r in manager.rectangles] == expected_str, f"Test 3 Failed:\nGot: {str([str(r) for r in manager.rectangles])}\nExpected: {expected_str}"
```

<!-- grade_id: cell-7bb962730a7db044, points: 1, grade: true, solution: false -->
```python
manager = RectangleManager()
manager.add_rectangle((1, 1), (4, 5))
manager.add_rectangle((0, 0), (2, 2))
manager.add_rectangle((3, 3), (5, 7))
manager.delete_rectangles_area_gt(8)
expected_str = (
    "[(0, 0), (2, 2)] 4 8\n"
    "[(3, 3), (5, 7)] 8 12"
)
assert str(manager) == expected_str, f"Test 4 Failed: Manager = {str(manager)}"
```

<!-- grade_id: cell-fcd7a370478eaa6d, points: 1, grade: true, solution: false -->
```python
try:
    rect = Rectangle((5, 5), (1, 1))  # Invalid: bs is not the bottom-left corner
    assert False, "Exception should be raised when bs > ad."
except ValueError as e:
    assert True, "Exception message should be raised."
```

---
### Exercise 3

Write a function `assign_to_category` that organizes a list of ordered dishes into categories based on a predefined list of dish-category pairs.

Function `assign_to_category`:
- takes two arguments:
   - `ordered_dishes`: A list of dishes ordered by a restaurant table, where each dish is a `string`.
   - `dish_categories`: A list of tuples where each tuple represents a dish and its corresponding category (e.g., `("Carbonara", "First Course")`).
- returns:
   - An empty dictionary if the `ordered_dishes` list is empty.
   - Otherwise, a dictionary where:
     - Keys are categories (e.g., `"First Course"`, `"Salads"`).
     - Values are lists of ordered dishes belonging to each category.

**NOTE:** If there is no category for an ordered dish, just do not include that dish in the output.

Write your code here:

```python
### Write your code here
def assign_to_category(ordered_dishes, dish_categories):
    result = {}
    if len(ordered_dishes) == 0:
        return result
    for dish in ordered_dishes:
        for category in dish_categories:
            if dish == category[0]:
                if category[1] in result:
                    result[category[1]].append(dish)
                else:
                    result[category[1]] = [dish]
    return result
```

To test your code run these cells:

<!-- grade_id: cell-2209dc856336d66e, points: 1, grade: true, solution: false -->
```python
# Test: Empty ordered dishes
ordered_dishes = []
dish_categories = [
    ("Carbonara", "First Course"),
    ("Caesar Salad", "Salads"),
]
expected = {}
result = assign_to_category(ordered_dishes, dish_categories)
assert result == expected, f"Test Failed: Obtained {result}, Expected {expected}"
```

<!-- grade_id: cell-9e55268423a2e0da, points: 1, grade: true, solution: false -->
```python
# Test: Ordered dishes with no matching categories
ordered_dishes = ["Sushi", "Burger"]
dish_categories = [
    ("Carbonara", "First Course"),
    ("Caesar Salad", "Salads"),
]
expected = {}
result = assign_to_category(ordered_dishes, dish_categories)
assert result == expected, f"Test Failed: Obtained {result}, Expected {expected}"
```

<!-- grade_id: cell-c1df7b43b1c94545, points: 1, grade: true, solution: false -->
```python
# Test: Basic functionality
ordered_dishes = ["Carbonara", "Caesar Salad", "Tiramisu"]
dish_categories = [
    ("Carbonara", "First Course"),
    ("Caesar Salad", "Salads"),
    ("Tiramisu", "Desserts"),
]
expected = {
    "First Course": ["Carbonara"],
    "Salads": ["Caesar Salad"],
    "Desserts": ["Tiramisu"],
}
result = assign_to_category(ordered_dishes, dish_categories)
assert result == expected, f"Test Failed: Obtained {result}, Expected {expected}"
```

<!-- grade_id: cell-7b138018caa90770, points: 2, grade: true, solution: false -->
```python
# Test: Multiple dishes in the same category
ordered_dishes = ["Carbonara", "Lasagna", "Tiramisu"]
dish_categories = [
    ("Carbonara", "First Course"),
    ("Lasagna", "First Course"),
    ("Tiramisu", "Desserts"),
]
expected = {
    "First Course": ["Carbonara", "Lasagna"],
    "Desserts": ["Tiramisu"],
}
result = assign_to_category(ordered_dishes, dish_categories)
assert result == expected, f"Test Failed: Obtained {result}, Expected {expected}"
```

<!-- grade_id: cell-20cdec6665634c46, points: 1, grade: true, solution: false -->
```python
# Test: Duplicate dishes in the ordered list
ordered_dishes = ["Carbonara", "Carbonara", "Tiramisu"]
dish_categories = [
    ("Carbonara", "First Course"),
    ("Tiramisu", "Desserts"),
]
expected = {
    "First Course": ["Carbonara", "Carbonara"],
    "Desserts": ["Tiramisu"],
}
result = assign_to_category(ordered_dishes, dish_categories)
assert result == expected, f"Test 6 Failed: Obtained {result}, Expected {expected}"
```